[1613571603] SERVICE ALERT: localhost;Swap Usage;CRITICAL;SOFT;1;SWAP CRITICAL - 0% free (0 MB out of 0 MB) - Swap is either disabled, not present, or of zero size.
[1613571663] SERVICE ALERT: localhost;Swap Usage;CRITICAL;SOFT;2;SWAP CRITICAL - 0% free (0 MB out of 0 MB) - Swap is either disabled, not present, or of zero size.
[1613571685] SERVICE ALERT: localhost;HTTP;WARNING;SOFT;1;HTTP WARNING: HTTP/1.1 401 Unauthorized - 691 bytes in 0.001 second response time
[1613571745] SERVICE ALERT: localhost;HTTP;WARNING;SOFT;2;HTTP WARNING: HTTP/1.1 401 Unauthorized - 691 bytes in 0.001 second response time
[1613571794] SERVICE ALERT: localhost;Root Partition;WARNING;SOFT;1;DISK WARNING - free space: / 9844 MiB (13.66% inode=56%):
[1613571805] SERVICE ALERT: localhost;HTTP;WARNING;SOFT;3;HTTP WARNING: HTTP/1.1 401 Unauthorized - 691 bytes in 0.001 second response time
[1613571822] SERVICE ALERT: localhost;SSH;CRITICAL;SOFT;1;connect to address 127.0.0.1 and port 22: Connection refused
[1613571828] SERVICE ALERT: localhost;Swap Usage;CRITICAL;SOFT;3;SWAP CRITICAL - 0% free (0 MB out of 0 MB) - Swap is either disabled, not present, or of zero size.
[1613571854] SERVICE ALERT: localhost;Root Partition;WARNING;SOFT;2;DISK WARNING - free space: / 9841 MiB (13.65% inode=56%):
[1613571865] SERVICE ALERT: localhost;HTTP;WARNING;HARD;4;HTTP WARNING: HTTP/1.1 401 Unauthorized - 691 bytes in 0.001 second response time
[1613571882] SERVICE ALERT: localhost;SSH;CRITICAL;SOFT;2;connect to address 127.0.0.1 and port 22: Connection refused
[1613571883] SERVICE ALERT: viewnext109;Jenkins;WARNING;SOFT;1;HTTP WARNING: HTTP/1.1 403 Forbidden - 1531 bytes in 0.004 second response time
[1613571888] SERVICE ALERT: localhost;Swap Usage;CRITICAL;HARD;4;SWAP CRITICAL - 0% free (0 MB out of 0 MB) - Swap is either disabled, not present, or of zero size.
[1613571914] SERVICE ALERT: localhost;Root Partition;WARNING;SOFT;3;DISK WARNING - free space: / 9841 MiB (13.65% inode=56%):
[1613571942] SERVICE ALERT: localhost;SSH;CRITICAL;SOFT;3;connect to address 127.0.0.1 and port 22: Connection refused
[1613571943] SERVICE ALERT: viewnext109;Jenkins;WARNING;SOFT;2;HTTP WARNING: HTTP/1.1 403 Forbidden - 1531 bytes in 0.003 second response time
[1613571974] SERVICE ALERT: localhost;Root Partition;WARNING;HARD;4;DISK WARNING - free space: / 9837 MiB (13.65% inode=56%):
[1613572002] SERVICE ALERT: localhost;SSH;CRITICAL;HARD;4;connect to address 127.0.0.1 and port 22: Connection refused
[1613572003] SERVICE ALERT: viewnext109;Jenkins;WARNING;SOFT;3;HTTP WARNING: HTTP/1.1 403 Forbidden - 1530 bytes in 0.004 second response time
[1613572063] SERVICE ALERT: viewnext109;Jenkins;WARNING;HARD;4;HTTP WARNING: HTTP/1.1 403 Forbidden - 1531 bytes in 0.002 second response time
[1613573259] SERVICE ALERT: viewnext109;Jenkins;WARNING;SOFT;1;HTTP WARNING: HTTP/1.1 403 Forbidden - 1532 bytes in 0.003 second response time
[1613573319] SERVICE ALERT: viewnext109;Jenkins;WARNING;SOFT;2;HTTP WARNING: HTTP/1.1 403 Forbidden - 1531 bytes in 0.002 second response time
[1613573431] SERVICE ALERT: viewnext109;Jenkins;OK;SOFT;3;HTTP OK: HTTP/1.1 200 OK - 61300 bytes in 0.439 second response time
[1613575745] HOST ALERT: viewnext109.sinensia.pw;DOWN;SOFT;1;CRITICAL - Could not interpret output from ping command
[1613575788] HOST ALERT: viewnext109.sinensia.pw;DOWN;SOFT;2;CRITICAL - Could not interpret output from ping command
[1613575840] SERVICE ALERT: viewnext109.sinensia.pw;HTTP;WARNING;HARD;1;HTTP WARNING: HTTP/1.1 400 Bad Request - 335 bytes in 0.083 second response time
[1613575848] HOST ALERT: viewnext109.sinensia.pw;DOWN;SOFT;3;CRITICAL - Could not interpret output from ping command
[1613575908] HOST ALERT: viewnext109.sinensia.pw;DOWN;SOFT;4;CRITICAL - Could not interpret output from ping command
[1613575936] HOST ALERT: viewnext109.sinensia.pw;DOWN;SOFT;5;CRITICAL - Could not interpret output from ping command
[1613575996] HOST ALERT: viewnext109.sinensia.pw;DOWN;SOFT;6;CRITICAL - Could not interpret output from ping command
[1613576021] HOST ALERT: viewnext109.sinensia.pw;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.15 ms
[1613576440] SERVICE ALERT: viewnext109.sinensia.pw;HTTP;OK;HARD;3;HTTP OK: HTTP/1.1 200 OK - 670 bytes in 0.001 second response time
[1613576536] SERVICE ALERT: viewnext109.sinensia.pw;HTTPS;CRITICAL;SOFT;1;connect to address 10.250.9.2 and port 443: Connection refused
[1613576656] SERVICE ALERT: viewnext109.sinensia.pw;HTTPS;CRITICAL;SOFT;2;connect to address 10.250.9.2 and port 443: Connection refused
[1613576776] SERVICE ALERT: viewnext109.sinensia.pw;HTTPS;CRITICAL;HARD;3;connect to address 10.250.9.2 and port 443: Connection refused